Das beliebteste Einstiegsbeispiel der Informatik ist die Ausgabe eines "Hello world!" auf der Kommandozeile. In der nachfolgenden Zelle ist dieses kleine Programm für sie schon implementiert. Sie müssen nur noch die Zelle auswählen und auf **RUN** drücken. Voila sie haben ihr erstes kleines Programm ausgeführt. Was müssen sie nun ändern, dass nicht mehr "Bitte ändere mich!", sondern "Hello world!" ausgegeben wird?

```python
print("Bitte ändere mich!")
```

Output:
```
Bitte ändere mich!
```

Nach diesem kleinen ersten Erfolg ist es Zeit ein Grundkonzept der Informatik kennenzulernen: Variablen! Das kommt ihnen wahrscheinlich aus dem Mathematikunterricht bekannt vor. Anstatt fester Werte werden Variablen übergeben, die beliebige Werte annehmen können. Variablen in Python folgen einer einfach Logik. **Variablenname = Variablenwert**. Wenn wir der Variable pi den Wert 3.1416 zuweisen wollen würde das wie folgt aussehen:

pi = 3.1416

**Aufgabe 1:**
In unserem nächsten Beispiel wollen wir mit Hilfe von Variablen dieselbe Ausgabe (Wir erinnern uns an "Hello world!") erzeugen. Hier können sie auch gleich ihren ersten Punkt abholen.

<!-- grade_id: cell-4f6cd1f1f0fe6c0c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
    #TODO: your code
variable = ""
### END SOLUTION
print(variable)
```

<!-- grade_id: cell-236366c29f7cfeac, points: 1, grade: true, solution: false -->
```python
assert variable == "Hello world!"
```

Was ist wenn wir unser "Hello world!" nun aber öfter nutzen wollen und es aber nicht jedes mal ausschreiben wollen? Genau für diesen Fall gibt es Methoden. Methoden erleichtern den Programmierer Altag ungemein, da sie erstens erlauben bereits geschriebenen Code wieder zu benutzen und zweitens den Code auch besser lesbar zu machen. Methoden in Python folgen dem Muster **def methoden_name():**. Alles was daraufhin in den nächsten Zeilen eingerückt geschrieben wird gehört zu dieser Methode. Die Methode könnt ihr dann einfach ausführen indem ihr methoden_name() schreibt. Zum Beispiel ist auch das print das ihr für "Hello world!" benutzt habt eine dieser Methoden. Die hat aber freundlicherweise schon mal jemand für euch bereitgestellt. In der nächsten Zeile ist das "Hello world!" Programm als Methode implementiert und ihr könnt es durch drücken auf play mal ausführen.

```python
#Methode die "Hello world!" ausgiebt.
def print_hello_world():
    print("Hello world!")
```

```python
print_hello_world()
```

Output:
```
Hello world!
```

Das ist natürlich nur ein einfaches Beispiel und kann noch beliebig komplex werden. Ihr könnt der Methode auch noch Werte übergeben, die ihr dann als Variablen innerhalb der Methode nutzen könnt. Diese Variablen nennt man dann Übergabeparameter und werden bei der Methodendefinition folgendermaßen deklariert: **def methoden_name( Übergabeparameter)**. Bei mehr als eienm Übergabeparameter teilt ihr diese einfach durch Kommata. Beispielsweise könntet ihr "Hello world!" als Übergabeparameter für eure Methode nutzen, um eure print Methode für alle Möglichen Eingaben zu erweitern und so genereller zu gestallten. Schaut euch das nachfolgende Beispiel an und spielt gerne ein bisschen damit herum. Ihr werdet das Gelernte für deie Abschlussfrage brauchen.

```python
#Methode die einen Satz druckt der aus den beiden Teilen satzteil1 und satzteil2 besteht.
def print_erweitert(satzteil1, satzteil2):
    print(satzteil1 + " " + satzteil2)
```

```python
print_erweitert("Hello", "world!")
```

Output:
```
Hello world!
```

Als finale Lektion für dieses Notebook lernt ihr noch wie eure Methode ein Ergebniss zurückgeben kann. Dies geschieht mit einem return statement innerhalb der Methode. Schaut euch das nachfolgende Beispiel gut an und Versucht daraufhin die Aufgabe 2 zu lösen.

```python
#Methode die zwei Satzteile, nämlich satzteil1 und satzteil2 zu einem Satz zusammenfügt,
#der durch ein Leerzeichen getrennt wird
def zusammenfuegen(satzteil1, satzteil2):
    return satzteil1 + " " + satzteil2
```

```python
zusammengefuegterSatz = zusammenfuegen("Hello", "world!")
print(zusammengefuegterSatz)
```

Output:
```
Hello world!
```

**Aufgabe 2:** Als kleines Beispiel fügen sie hier bitte den Code für eine einfache Addition ein. Zum Beispiel soll als Ergebnis von add(1,1), also 1+1, 2 zurück gegeben werden.

<!-- grade_id: cell-200957a72ade9107, points: 0, grade: false, solution: true -->
```python
def add(x, y):
    ### BEGIN SOLUTION
    #TODO: your code
    ### END SOLUTION
```

Output:
```
File "<ipython-input-2-9f1f3005cfed>", line 4
    ### END SOLUTION
                    ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: cell-c2054285366131a4, points: 3, grade: true, solution: false -->
```python
assert add(1, 1) == 2
assert add(1, 0) == 1
assert add(1, -1) == 0
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-add3a6afb776> in <module>
----> 1 assert add(1, 1) == 2
      2 assert add(1, 0) == 1
      3 assert add(1, -1) == 0

NameError: name 'add' is not defined
```

```python

```